![Banner](Banner.jpg)

# Definici贸n de un MDP para Gridworld

En este tutorial vamos a definir un ambiente de prueba com煤nmente utilizado en el aprendizaje por refuerzo; el ambiente de Gridworld. Vamos a definir el ambiente de Gridworld como un Markov Decision Process (MDP) para luego establecer soluci贸n con los diferentes m茅todos basados en modelos y libres de modelos.
Gridworld es un ambiente cl谩sico de prueba dentro del aprendizaje por refuerzo. Durante este tutorial definiremos el modelo b谩sico del ambiente, que extenderemos incrementalmente de acuerdo con las necesidades de los algoritmos espec铆ficos de soluci贸n que desarrollaremos a lo largo del curso.

Los ambientes de aprendizaje por refuerzo se pueden definir en dos (ambiente, agente) o tres (ambiente, agente, aprendizaje) m贸dulos. Por simplicidad, dentro de nuestra implementaci贸n de Gridworld (e implementaciones posteriores) utilizaremos dos m贸dulos. Sin embargo, cabe notar que una implementaci贸n utilizando tres m贸dulos ofrece una mejor divisi贸n de las caracter铆sticas y por lo tanto una implementaci贸n m谩s f谩cil de mantener y modificar.

Para cada uno de los m贸dulos a continuaci贸n daremos una clase principal de c贸digo sobre la cual se deben hacer todas las modificaciones. Cada una de las funciones a implementar se presentar谩n progresivamente. Para cada una de las funciones a implementar se dar谩 un conjunto (m铆nimo) de prueba.


## Ambiente 

El ambiente de Gridworld se define como una cuadricula de `nxm`. El ambiente tiene obst谩culos, es decir casillas por las cuales no puede pasar el agente. Al chocar con un obst谩culo, el agente termina en la misma casilla en la que estaba. Un comportamiento similar se observa cuando el agente trata de salir de los bordes del ambiente. Adem谩s, el ambiente tiene una casilla de inicio definida, y algunas casillas de salida (determinadas por la recompensa asociada a ellas). Un ejemplo del ambiente para el caso `3x4` se muestra a continuaci贸n.

![gridworld.png](gridworld.png)

En este ejemplo del ambiente la casilla de inicio es la casilla inferior izquierda y tiene como objetivo llegar a la casilla de salida con recompensa 1. La otra casilla de salida tiene recompensa -1.


#### 驴C贸mo podemos codificar el ambiente?

La definici贸n del ambiente de Gridworld esta definida por:
1. Una cuadr铆cula rectangular (`grid`), con dimensiones `(n,m)` dadas por par谩metro, donde la casilla superior izquierda esta en la posici贸n (0,0). Definiremos las casillas por las que puede pasar el agente como espacios en blanco y las casillas por las que no puede pasar el agente con un `'#'`.
2. Las recompensas de cada casilla de la cuadr铆cula estan definidas dentro de la definici贸n de `grid`.
    - +1 para la casilla objetivo 
    - -1 para las casillas de trampa
3. Un atributo con el estado actual (`state`) en el que se encuentra el agente. Por defecto este estado ser谩 la posici贸n marcada como `S`.

Un ejemplo de una cuadr铆cula de 3x4, como se mostr贸 anteriormente, ser铆a as铆:

 ```python  
    board = [[' ', ' ', ' ',  +1],
            [' ', '#', ' ',  -1],
             ['S', ' ', ' ', ' ']] 
```

La definici贸n del ambiente un ejemplo de recompensas se ver铆a as铆: 

```python
    def __init__(self, board):
        # layout
        self.grid = copy_elements_from_the_board
```


#### 1. Estructura del ambiente

Defina la clase `Gridworld`, que recibe una cuadr铆cula con la descripci贸n del tablero `board`, definidos como en el ejemplo. La informaci贸n codificada en la cuadr铆cula ser谩 almacenada en el atributo `grid` del ambiente (este atributo corresponde al mdp donde se almacena la informaci贸n del ambiente, la funci贸n de transici贸n y las recompensas). Adicionalmente, para facilitar el uso del mdp, definimos los atributos para guardar la informaci贸n de la cuadr铆cula; las filas (`nrows`) y columnas (`ncols`). esta informaci贸n se da por par谩metro (`dimensions`), al instanciar la clase, como una tupla con los valores respectivos. Adicionalmente la clase debe tener los atributos `initial_state` y `state` que corresponden al estado inicial y el estado actual del agente en el ambiente, respectivamente. Estos atributos se guardan como tuplas.

Finalmente, el atributo `grid` se debe codificar de tal forma que las casillas prohibidas (marcadas como `'#'`) no deben tener ning煤n valor asignado (su valor debe ser `None`), las casillas en blanco deben tener valor `0`, y las casillas con recompensas asociadas deben tener el valor de la recompensa como su valor. Note que con esta codificaci贸n de `grid`, y teniendo en cuenta que las acciones son determin铆sticas, estamos codificando directamente las recompensas del ambiente. Si este no fuera el caso, ser铆a necesario definir un nuevo atributo `rewards`.

<!-- grade_id: cell-11f1d5794939d3dd, points: 0, grade: false, solution: true -->
```python
import matplotlib.pyplot as plt
import matplotlib.patches as patches

class Gridworld:
    def __init__(self, board: list[list[str]], dimensions: tuple[int,int]):
        self.nrows, self.ncols = dimensions
        self.grid = [[None for _ in range(self.ncols)] for _ in range(self.nrows)]
        self.initial_state = None
        self.state = None
        self.board = board

        for i in range(self.nrows):
            for j in range(self.ncols):
                cell = board[i][j]

                if cell == '#':
                    self.grid[i][j] = None
                elif cell == ' ':
                    self.grid[i][j] = 0
                elif cell == 'S':
                    self.initial_state = (i, j)
                    self.state = (i, j)
                    self.grid[i][j] = 0
                else:
                    try:
                        self.grid[i][j] = int(cell)
                    except:
                        self.grid[i][j] = cell

    def get_current_state(self) -> tuple[int,int]:
        return self.state

    def get_possible_actions(self, state:tuple[int,int]) -> list[str]:
        pass

    def do_action(self, action: str) -> tuple[float, tuple[int,int]]:
        pass

    def reset(self):
        self.state = self.initial_state
        return self.state

    def is_terminal(self) -> bool:
        i, j = self.state
        value = self.grid[i][j]
        return value == 1 or value == -1

    def get_action_index(self, action:str) -> int:
        actions = ['up', 'right', 'down', 'left']
        for idx, a in enumerate(actions):
            if a == action:
                return idx

    def plot(self):
        fig1 = plt.figure(figsize=(10, 10))
        ax1 = fig1.add_subplot(111, aspect='equal')

        for i in range(0, len(self.grid)+1):
            ax1.axhline(i , linewidth=2, color="#2D2D33")
            ax1.axvline(i , linewidth=2, color="#2D2D33")

        (i,j)  = self.initial_state
        ax1.add_patch(patches.Rectangle((j, self.nrows - i -1), 1, 1, facecolor = "#F6D924"))

        for j in range(len(self.grid[0])):
            for i in range(len(self.grid)):
                if self.grid[i][j] == 1:
                    ax1.add_patch(patches.Rectangle((j,self.nrows - i -1), 1, 1, facecolor = "#68FF33"))
                if self.grid[i][j] == None:
                    ax1.add_patch(patches.Rectangle((j,self.nrows - i -1), 1, 1, facecolor = "#6c7780"))
                if self.grid[i][j] == -1:
                    ax1.add_patch(patches.Rectangle((j,self.nrows - i -1), 1, 1, facecolor = "#cc0000"))

        plt.scatter(self.state[1] + 0.5, self.nrows - self.state[0] - 1 + 0.5, 
                    s=100, color="blue", marker="o", zorder=10)
        
        for i in range(len(self.grid)):
            for j in range(len(self.grid[0])):
                if self.grid[i][j] is not None:
                    ax1.text(j+0.5, self.nrows-i-1+0.5, self.grid[i][j], ha='center', va='center')

        plt.axis("off")
        plt.show()

```

<!-- grade_id: cell-d02ebef01da3ead9, points: 0, grade: true, solution: false -->
```python
#Pruebas estructura del ambiente
### BEGIN TESTS
board = [[' ', ' ', '-1'],
      ['S', '#', ' '],
      [' ', ' ', '+1']]

grid = Gridworld(board, (3,3))

try:
    grid.nrows
except:
    print("El atributo nrows no est谩 definido")
try:
    grid.ncols
except:
    print("El atributo ncols no est谩 definido")
try:
    grid.initial_state
except:
    print("El atributo initial_state no est谩 definido")

try:
    grid.state
except:
    print("El atributo state no est谩 definido")

try:
    grid.grid
except:
    print("El atributo grid no est谩 definido")

assert grid.ncols == 3, "El valor de las columnas debe ser el dado por par谩metro en el atributo dimensions"
assert grid.nrows == 3, "El valor de las columnas debe ser el dado por par谩metro en el atributo dimensions"
assert grid.grid[1][1] == None, "El valor de las casillas sobre las que no puede pasar el agente debe ser None"
assert grid.grid[0][2] == -1, "El valor de las casillas con recompensa debe ser el valor num茅rico (tipo int)"
assert grid.initial_state == (1,0), "La posici贸n inicial del agente debe ser la definida en el estado S de la cuadr铆cula dada"
assert grid.initial_state == grid.state, "La posici贸n del agente debe ser igual a su posici贸n inicial si ning煤n paso se ha dado"
### END TESTS
```

```python
#Pruebas adicionales

```

#### 2. Movimientos del agente

Defina la funci贸n `do_action` que ejecuta la acci贸n tomada por el agente dentro de la cuadr铆cula. Esta funci贸n recibe como par谩metro la acci贸n a ejecutar y retorna el valor de la casilla de llegada de la acci贸n y el estado de llegada de la acci贸n (como una tupla). Note que los movimientos fuera del tablero o las casillas prohibidas no deber铆an tener ning煤n efecto en la posici贸n del agente (el agente se debe mantener en la misma posici贸n de partida).

En esta versi贸n de gridworld vamos a trabajar con acciones determin铆sticas, es decir el movimiento deseado del agente (`up`,`right`,`down`,`left`) siempre resultara en el estado esperado. 


```python
class Gridworld(Gridworld):
    def do_action(self, action: str) -> tuple[float, tuple[int,int]]:
        i, j = self.state
        next_i, next_j = i, j

        if action == 'up':
            next_i -= 1
        elif action == 'down':
            next_i += 1
        elif action == 'left':
            next_j -= 1
        elif action == 'right':
            next_j += 1

        # Validar l铆mites y obst谩culos
        if next_i < 0 or next_i >= self.nrows or next_j < 0 or next_j >= self.ncols:
            next_i, next_j = i, j
        elif self.grid[next_i][next_j] is None:
            next_i, next_j = i, j

        self.state = (next_i, next_j)
        reward = self.grid[next_i][next_j]

        return reward, self.state
        
```

<!-- grade_id: cell-4c2ec1213aac0a9d, points: 0, grade: true, solution: false -->
```python
#Pruebas de movimiento
### BEGIN TESTS
board = [[' ', ' ', '-1'],
      ['S', '#', ' '],
      [' ', ' ', '+1']]

grid = Gridworld(board, (3,3))

try:
    grid.do_action
except:
    print("La funci贸n do_action no est谩 definida")

val, state = grid.do_action('up')
assert state == (0,0), f"El movimiento hacia arriba no esta modificando correctamente al agente, se movi贸 a {state} y deber铆a ser (0,0)"
val, state = grid.do_action('right')
assert state == (0,1), f"El movimiento hacia la derecha no esta modificando correctamente al agente, se movi贸 a {state} y deber铆a ser (0,1)"
val, state = grid.do_action('right')
assert val == -1, f"El valor de las casillas no se esta retornando correctamente"
val, state = grid.do_action('down')
assert state == (1,2), f"El movimiento hacia la derecha no esta modificando correctamente al agente, se movi贸 a {state} y deber铆a ser (1,2)"
val, state = grid.do_action('left')
assert state == (1,2), f"El movimiento hacia la derecha no esta modificando correctamente al agente, se movi贸 a {state} y deber铆a ser (1,2) porque hay un obst谩culo en (1,1)"

### END TESTS
```

```python
#Pruebas adicionales


```

#### 3. Estado actual

Defina la funci贸n `get_current_state` que retorna el estado actual del agente (la casilla donde se encuentra el agente). Esta funci贸n no recibe ning煤n par谩metro y retorna el estado actual como una tupla.

<!-- grade_id: cell-44709443d150a086, points: 0, grade: true, solution: false -->
```python
#Pruebas estado actual
### BEGIN TESTS
board = [[' ', ' ', '-1'],
      ['S', '#', ' '],
      [' ', ' ', '+1']]

grid = Gridworld(board, (3,3))

try:
    grid.get_current_state
except:
    print("La funci贸n get_current_state no est谩 definida")

assert grid.get_current_state() == (1,0), "La posici贸n del agente debe ser la posici贸n inicial si no se realiza ning煤n movimeinto"
grid.do_action('up')
assert grid.get_current_state() == (0,0), f"La posici贸n del agente no se esta modificando correctamente, recibi贸 {grid.get_current_state()} cuando deber铆a ser (0,0), revise los movimientos"
grid.state = (1,0)
grid.do_action('right')
assert grid.get_current_state() == (1,0), "El agente no se mantiene en la misma posici贸n cuando trata de moverse a la casilla prohibida"
### END TESTS
```

```python
#Pruebas adicionales


```

#### 4. Obtener las acciones

Defina la funci贸n `get_possible_actions` que recibe el estado actual del agente por par谩metro y retorna una lista de las acciones v谩lidas para el estado dado.

Tenga en cuenta que en esta versi贸n de Gridworld, todas las acciones (i.e., los movimientos en las cuatro direcciones) son posibles para el agente en cada una de las casillas regulares de la cuadr铆cula. Las casillas de salida tienen 煤nica acci贸n posible `'exit'` y las casillas prohibidas no tienen ning煤na acci贸n asociada (una lista vac铆a de acciones).

```python
class Gridworld(Gridworld):
    def get_possible_actions(self, state):
        i, j = state
        cell = self.board[i][j]

        # Casilla bloqueada
        if cell == '#':
            return []

        # Casilla terminal
        if cell in ['-1', '+1']:
            return ['exit']

        # Casilla normal o inicial
        return ['up', 'right', 'down', 'left']
```

<!-- grade_id: cell-b829a305b1615d60, points: 0, grade: true, solution: false -->
```python
#Pruebas de acciones posibles
### BEGIN TESTS
board = [[' ', ' ', '-1'],
      ['S', '#', ' '],
      [' ', ' ', '+1']]

grid = Gridworld(board, (3,3))

try:
    grid.get_possible_actions
except:
    print("La funci贸n get_possible_actions no est谩 definida")

a1 = grid.get_possible_actions(grid.state)
a2 = grid.get_possible_actions((1,0))
a3 = grid.get_possible_actions((1,2))
a4 = grid.get_possible_actions((0,1))
a5 = grid.get_possible_actions((2,1))
ap = grid.get_possible_actions((1,1))
ar1 = grid.get_possible_actions((0,2))
ar2 = grid.get_possible_actions((2,2))
assert ap == [], "Las acciones para las casillas prohibidas deben estar vacias"
assert a1 == a2 == a4 == a5, f"Todas las casillas regulares deben tener el mismo conjunto de acciones {['up', 'right', 'down', 'left']}"
assert ar1 == ar2, f"Todas las casillas de salida deben tener el mismo conjunto de acciones {['exit']}"
### END TESTS
```

```python
#Pruebas adicionales


```


#### 5. Reinicializar el ambiente

Defina la funci贸n `reset` que no recibe par谩metros ni retorna ning煤n valor. El efecto de esta funci贸n es restablecer el ambiente a su estado inicial.



```python

```

<!-- grade_id: cell-365f971ddb690a05, points: 0, grade: true, solution: false -->
```python
#Pruebas de reset
### BEGIN TESTS
board = [[' ', ' ', '-1'],
      ['S', '#', ' '],
      [' ', ' ', '+1']]

grid = Gridworld(board, (3,3))

try:
    grid.reset
except:
    print("La funci贸n reset no est谩 definida")

grid.state = (2,2)
grid.reset()
assert grid.get_current_state() == grid.initial_state, f"No se esta retornando al estado inicial. {grid.get_current_state()} y no {grid.initial_state}"
grid.state = (1,1)
grid.reset()
assert grid.get_current_state() == grid.initial_state, f"No se esta retornando al estado inicial. {grid.get_current_state()} y no {grid.initial_state}"
### END TESTS
```

```python
#Pruebas adicionales


```


#### 6. Estados terminales

Defina la funci贸n `is_terminal` que determina si el agente est谩 en un estado final o no. En nuestro caso los estados finales o los estados de salida estar谩n determinados por las casillas con recompensa 1 o -1.
Esta funci贸n no recibe par谩metros y retorna un booleano determinando si el agente est谩 en un estado final o no.

<!-- grade_id: cell-b5125fef529b7d6a, points: 0, grade: true, solution: false -->
```python
#Pruebas terminaci贸n
### BEGIN TESTS
board = [[' ', ' ', '-1'],
      ['S', '#', ' '],
      [' ', ' ', '+1']]

grid = Gridworld(board, (3,3))

try:
    grid.is_terminal
except:
    print("La funci贸n is_terminal no est谩 definida")

grid.state = (2,2)
assert grid.is_terminal() == True, f"No se identifican correctamente los estados de salida como finales"
grid.state = (0,2)
assert grid.is_terminal() == True, f"No se identifican correctamente los estados de salida como finales"

for i in range(grid.nrows):
    for j in range(grid.ncols):
        grid.state = (i,j)
        if (i,j) == (0,2) or (i,j) == (2,2):
            assert grid.is_terminal() == True, f"No se identifican correctamente los estados de salida como finales"
        else:
            assert grid.is_terminal() == False, f"La casilla {(i,j)} se estan identificando como estado final"

### END TESTS
```

```python
#Pruebas adicionales


```

## Preguntas de reflexi贸n

Considere las siguientes situaciones y piense cual deber铆a ser el resultado del MDP.

1. Plantee el problema de MDP para cada una de las casillas. Especifique el estado de inicio, las transiciones y su probabilidad. Para cada acci贸n suponga una probabilidad de 茅xito de 0.25 (el 0.75 restante se resuelve uniformemente entre las acciones restantes).
驴C贸mo ser铆an las recompensas esperadas para cada estado?

2. Bajo la definci贸n del problema anterior, suponga que ahora cada acci贸n tiene una probabilidad de 茅xito de 60%, con probabilidad de 30% que se ejecutar谩 la sigiente acci贸n (en direcci贸n de las manesillas del reloj) y con probabilidad de 10% que no se ejecute ninguna acci贸n (el agente se queda quieto). Bajo estas condiciones, 驴C贸mo ser铆an las recompensas esperadas para cada estado? 


```python

```